<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 73723  -->
  <question type="stack">
    <name>
      <text>Ellipse, polar form (from HELM: 17.2.3.10 ellipsis vertices)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Locate the (Cartesian) coordinates for the vertices of \(r=\dfrac{{@cc*e*d@}}{{@cc@}-{@cc*e@}\cos\theta}\)</p>
<p>\(v_1 = (\)[[input:ans1]]\(,\ \)[[input:ans2]]\()\) [[validation:ans1]][[validation:ans2]]</p>
<p>[[feedback:prt1]]</p>
<p>\(v_2 = (\)[[input:ans3]]\(,\ \)[[input:ans4]]\()\) [[validation:ans3]][[validation:ans4]]</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>We aim to find the Cartesian coordinates for the vertices of \(r=\dfrac{{@cc*e*d@}}{{@cc@}-{@cc*e@}\cos\theta}\).</p>
<p>We first note that the general equation for a conic in polar coordinates is \(r = \dfrac{ed}{1-e\cos\theta}\), where \(e\) is the eccentricity of the ellipse and \(-d\) is its directrix. Dividing the top and bottom by {@cc@} gives us \(r=\dfrac{{@cc*e*d@}}{{@cc@}-{@cc*e@}\cos\theta}\), and therefore we find the the eccentricity \(e = {@e@}\).</p>
<p>As the ellipse is centred on the x-axis, it's vertices will occur when \(\theta = 0\) or \(\theta = \pi\). Thus we can find the two values of r as:</p>
<p>For \(\theta=0\) \[ \begin{aligned} r &amp;= \dfrac{ed}{1-e\cos\theta}\\ &amp;= \dfrac{ed}{1-e}\\ &amp;= \frac{{@e*d@}}{1-{@e@}}\\ &amp;={@ta2x@} \\ &amp;\implies x = {@ta2x@} \end{aligned} \]</p>
<p>For \(\theta = \pi\) \[ \begin{aligned} r &amp;= \dfrac{ed}{1-e\cos\theta}\\ &amp;= \dfrac{ed}{1+e}\\ &amp;= \frac{{@e*d@}}{1+{@e@}}\\ &amp;={@e*d/(1+e)@}\\ &amp;\implies x = {@ta1x@} \end{aligned} \]</p>
<p>Finally, as we know that the vertices lie on the x-axis, and that, we can deduce that \(r=x\) at these points, and therefore conclude that our vertices are: \(A = ({@ta1x@},{@ta1y@}),\ B = ({@ta2x@},{@ta2y@}).\)&lt;.p&gt;</p>
<p>See the graph of this ellipsis below. The points \(C\) and \(D\) are the foci and C is at the begining of the axis.</p>
<p>[[jsxgraph]] console.log(Math.abs({#ta1x#})) var max_x = Math.max(Math.abs({#ta1x#}), Math.abs({#ta2x#})) var board = JXG.JSXGraph.initBoard(divid, { boundingbox: [-max_x, max_x, max_x, -max_x], keepaspectratio: true, showCopyright: false, axis: true, showNavigation: true, }); JXG.Options.text.useMathHax = true; var A = board.create('point', [{#ta1x#}, {#ta1y#}]); var B = board.create('point', [{#ta2x#}, {#ta2y#}]); var f1 = board.create('point', [0, 0]); var f2 = board.create('point', [{#ta2x+ta1x#}, {#ta2y#}]); var polar = board.create('curve', [function(theta){ return ({#cc*e*d#}/({#cc#}-({#cc*e#} * Math.cos(theta)))) }, [0, 0], 0, 2*Math.PI], {curveType: 'polar'}); [[/jsxgraph]]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026042200</text>
    </stackversion>
    <questionvariables>
      <text>/*The value of e needs to be bigger (or equal) than 1/2 for the ellipse not to look like circle*/

e: rand([1/2, 2/3, 4/5]); /*ecentricity*/
bb:rand(4)+2;
d:bb/e;  /*directrix*/
cc:rand(2)+2;

rr: cc*e*d/(cc-cc*e*cos(theta))

ta1x:-e*d/(1+e);
ta1y:0;
ta2x:e*d/(1-e);
ta2y:0;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(r=\dfrac{{@cc*e*d@}}{{@cc@}-{@cc*e@}\cos\theta}, \quad e={@e@}, \quad d= {@d@}, \quad v_1 = ({@ta1x@}, {@ta1y@}), \quad v_2 = ({@ta2x@}, {@ta2y@})\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p><i class="fa fa-check"></i> Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><i class="fa fa-adjust"></i> Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><i class="fa fa-times"></i> Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1x</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta1y</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta2x</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta2y</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1x</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta1y</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta2x</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2y</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>you repeated the same vertex twice</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta2x</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta2y</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta1x</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta1y</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1048750018</deployedseed>
    <deployedseed>1634482742</deployedseed>
    <deployedseed>968816142</deployedseed>
    <deployedseed>1115747617</deployedseed>
    <deployedseed>1276341706</deployedseed>
    <deployedseed>922819660</deployedseed>
    <deployedseed>1495835362</deployedseed>
    <deployedseed>1888373274</deployedseed>
    <deployedseed>1708154191</deployedseed>
    <deployedseed>1223470590</deployedseed>
    <deployedseed>1172497962</deployedseed>
    <deployedseed>1292703097</deployedseed>
    <deployedseed>555121882</deployedseed>
    <deployedseed>1453577834</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1x</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1y</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta2x</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta2y</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>ta2x</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2y</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta1x</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta1y</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-5-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>ta1x</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1y</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta1x</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta1y</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>