<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 1145  -->
  <question type="stack">
    <name>
      <text>Free-text with final answer (calculus)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<div class="stack-comment">
<h2>Free-text input</h2>This question asks students to type in their complete answer using free-text input.  Instructions on how to do this are below.  The final answer will automatically be extracted and assessed with a normal STACK question.  Free-text can also be used with manual grading, or combined with other input types, as needed.
</div>
<p>Let \(f'(x)= {@dq1@}\). Given that \(f({@x0@})={@f0@}\), find \(f(x)\).</p>
<p>Work line by line below, justifying your answer fully.  Your last line should be your answer \(f(x)=...\).</p>
<div class="free-text-container">
[[input:ans1]] [[validation:ans1]]
[[ascii input="ans1"]]
  [[extractor targetinput="ans2" type="lastexpr"/]]
[[/ascii]]
</div>
<p>[[hint title="Input help"]][[commonstring key="free_text_fact"/]][[/hint]]</p>
<p style="display:none">[[input:ans2]]</p>
<p>[[validation:ans2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>[[ascii]]{@ta1@}[[/ascii]]</text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026042402</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[x0:0;
u1:2*x^2+1
k1:4;
eq1:k1*sqrt(u1)+1;
f0:ev(eq1,x=x0);
dq1:diff(eq1,x);
ta1:castext("Let
`
u={#u1#} => (du)/(dx) = {#diff(u1,x)#}
`
Hence
`
f(x) = int {#dq1#} dx
     = int 2/sqrt(u) du.
     = {#k1#} sqrt(u)+c
     = {#k1#} sqrt({#u1#})+c
`
Using the given values `5 = 4sqrt({#ev(u1,x=x0)#})+c => c=1`, we have
`
f(x) = {#eq1#}
`");
ta2:f(x)=eq1;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt2]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[[[commonstring key="symbolicprtcorrectfeedback"/]] [[commonstring key="defaultprtcorrectfeedback"/]]]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[[[commonstring key="symbolicprtpartiallycorrectfeedback"/]] [[commonstring key="defaultprtpartiallycorrectfeedback"/]]]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[[[commonstring key="symbolicprtincorrectfeedback"/]] [[commonstring key="defaultprtincorrectfeedback"/]]]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>freetext</type>
      <tans>ta1</tans>
      <boxsize>80</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>monospace:true, manualgraded:false</options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>5</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa2:if equationp(ans2) then rhs(ans2) else ans2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>sa2</sans>
        <tans>eq1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>f'(x)</description>
        <answertest>AlgEquiv</answertest>
        <sans>diff(sa2,x)</sans>
        <tans>dq1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>The derivative of your answer is correct, well done.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>When we differentiate your answer we get \[ f'(x) = {@diff(sa2,x)@} \neq {@dq1@} \] so you must have done something wrong.</text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>f(0)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ev(sa2,x=x0)</sans>
        <tans>f0</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text>Your answer has \(f({@x0@})={@f0@}\), well done.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>When we evaluate your answer at \(x={@x0@}\) we get \[ f({@x0@}) = {@ev(sa2,x=x0)@} \neq {@f0@} \] so you must have done something wrong.</text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt2-1-T ]</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Doesn't write f(x)=</description>
      <testinput>
        <name>ans1</name>
        <value></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>eq1</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt2-1-T ]</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Wrong f(0)</description>
      <testinput>
        <name>ans1</name>
        <value></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>eq1+2</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>[ prt2-1-F | prt2-2-T | prt2-3-F ]</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Just wrong!</description>
      <testinput>
        <name>ans1</name>
        <value></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>x*eq1</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>[ prt2-1-F | prt2-2-F | prt2-3-F ]</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>